**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Change data types and add columns

## Import pandas

```python
import pandas as pd
```

## Import data

```python
# URL of data
URL = "https://raw.githubusercontent.com/kirenz/datasets/master/height_clean.csv"
df = pd.read_csv(URL)
```

```python
df.dtypes
```

## Change data type

- There are several methods to [change data types in pandas](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.astype.html):

The most common method to change the data type is:

- `.astype()`: Convert to a specific type (like "int32", "float" or "catgeory")
- `.astype(str)`: Convert to string  
  
More options:  
  
- `to_datetime`: Convert argument to datetime.
- `to_timedelta`: Convert argument to timedelta.
- `to_numeric`: Convert argument to a numeric type.
- `numpy.ndarray.astype`: Cast a numpy array to a specified type.

### Categorical data

- Categoricals are a pandas data type corresponding to categorical variables in statistics. 


- A categorical variable takes on a limited, and usually fixed, number of possible values (categories). Examples are gender, social class, blood type, country affiliation, observation time or rating via Likert scales.


- Convert variable "gender" to a category dtype:

```python
df["gender"] = df["gender"].astype("category")
```

### String data

- In our example, `id` is not a number (we can't perform calculations with it)
- It is just a unique identifier so we should transform it to a simple string (object)

Hint:

```python
df['___'] = df['___'].___(___)
```


<!-- grade_id: id, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_id, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.select_dtypes(include=["number"]).columns.to_list() == ['height', 'average_height_parents']
assert df.select_dtypes(include=["object"]).columns.to_list() == ['name', 'id']
```

## Add new columns

### Constant

- Add a new variable called "number" to df 
- The new variable should have the number 42 in all rows


Hint:
  
```python
df["___"] = ___
```

<!-- grade_id: add, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_add, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.loc[[0], "number"].to_list() == [42]
```

### From existing

 - Create new columns from existing columns

```python
# we use numpy to add some data with a nornmal distribution 
import numpy as np

# calculate height in m (from cm)
df['height_m'] = df.height/100 

# add 20 random normal distributed numbers with a mean of 45 and standard deviation of 5
df['weight'] = round(np.random.normal(45, 5, 20) * df['height_m'],2)

# calculate body mass index
df['bmi'] = round(df.weight / (df.height_m * df.height_m),2)
```

```python
df.head()
```

```python
# we need datetime to add a date
from datetime import datetime
```

```python
df["date"] = datetime.today().strftime('%Y-%m-%d')

df.head(3)
```